# Hacking the Humanities Week 1: First Steps in Python

Welcome to *Hacking the Humanities*. If this is your first time coding, and you haven't yet watched the welcome lecture on Moodle or completed the notebook in the [welcome](../welcome) folder of your CoCalc project, I would recommend you do both now.

Otherwise, let's get started with the course!

In this first week of the course, we are going to learn to define our own custom **functions** and **objects** in Python. By doing so, we will have the opportunity to learn all the basic components of the Python programming language, in particular *variables*, *functions*, *operators*, *packages* and *types* or *object classes*. These terms should already be familar to you if you have watched the explainer videos on Moodle.

## Section 1: Variables, Functions and Operators

A **variable** is like a pronoun (they, it, this, that) or a pronumeral (x, y). It is a placeholder to which you can assign a value. By convention, variable names are written in [snake_case](https://en.wikipedia.org/wiki/Snake_case). There cannot be a space in a variable name. You can assign a value to a variable using `=`.

If you just type a value without giving it a name, e.g. `7`, `"Star Wars"`, or `[1,2,3,4]`, this is called a **literal**.

Execute the cell below to see how to assign a value to a variable. To execute the cell, click on it, then press <kbd>Ctrl</kbd>/<kbd>Control</kbd> + <kbd>Enter</kbd>.

```python
my_var = 'This variable is a string'
print(my_var)
```

A **function** is a command you can use to make Python do something. Most functions require some *arguments* or *parameters* to be provided, which tell the computer how it should perform the function. You can think of a function as a *verb*, and the parameters as *nouns*.

An example: The `int()` function converts its input into an integer. Execute the cell below to see how it works.

```python
my_string = '100'
print(f'my_string is a {type(my_string)}')
print(my_string + ' billion dollars')

# Let's convert my_string into an integer using int()
my_int = int(my_string)
print(f'my_int is a {type(my_int)}')
print(my_int + 10)
```

In the code cell above, you will also have seen a `+`. This is an example of an **operator**. Operators are a bit like functions, in that they command the computer to do something. But whereas a function is a word with paraentheses, e.g. `do_something()`, an operator is simply a symbol or a word that in between two variables or literals you wish to affect. Common operators in Python include:

Symbol | Operation 
--- | --- 
`=` | Assignment
`==` | Equality
`+` | Addition or Concatenation
`-` | Subtraction
`*` | Multiplication
`/` | Division (with decimal point)
`//` | Division (whole number, drop the remainder)
`%` | Modulo (divide and just keep the remainder)

As you will have seen above, `+` does something different to strings and numbers:

```python
# `+` adds two numbers:
print(26.5 + 81.332)

# But it concatenates two strings
print('Arnold' + 'Schwarzenegger')
```

### Practice 1.1: Assignment vs. Equality

As mentioned in the explainer videos, there are two different kinds of equals sign in Python (as in many other programming languages). A single equals sign, `=`, is the *assignment operator*. It assigns a value to a variable. The double equals sign, `==`, is the *equality operator*. It checks if one variable or literal is equal to another. There are a range of similar operators in Python, such as `!=`, which is the *negation operator*, or `>=`, which is *greater than or equal to*, or `is` and `is not`, which are similar to `==` and `!=`, but check for *identity* rather than *equality*.

For this practice, experiment with the code cell below. To get you started, here are some questions you could consider:
* Does `'étrange'` equal `'etrange'`? How about `'Joe'` and `'joe'`?
* Does `7. == 7` yield the same result as `7. is 7`? How about `!=` and `is not`?
* Is `'27'` greater than or equal to `5`? (Note the quotation marks)
* What happens if you try to execute `82 = 82`? Why does that happen?

Python's different operators are explained in the [Expressions](https://docs.python.org/3/reference/expressions.html?highlight=operators#comparisons) chapter of the official documentation.

```python
my_var = 'A string'

my_var == 'a string'
```

### Assignment 1.2: Define your own function

In order to do real programming in Python, you need to be able to define your own functions. To define a function, you need to use a `def` statement. This is the format:

```
def function_name(parameter_1, parameter_2=45, parameter_3="default value", etc.):
    """A docstring describing what your function does"""

    # Some code which does something to the parameters
    # NB: These lines must be indented. By default you indent four spaces, which you can do using the 'tab' key

    output = parameter_1 // parameter_2 * 7

    # At the end of a function, you should usually have a `return` statement, which says what the output of the function will be.
    # If there is no 'return' statement, then the function will not have an output.

    return output
```

When you are developing code, it is a good practice to include comments explaining what it does, so that other people can understand your code—and so that you can understand it when you come back to it later. There are two ways of writing a comment in Python:

```
# A hashtag turns everything after it into a comment
```
or
```
"""
You can use triple quotation marks
to create a multiline comment.
Triple inverted commas (''') work as well.
"""
```

Now you know how to create a variable, use some basic operators and define a function, here is your first assignment: create a function called `concatenate()` that concatenates two strings. To do this:
* Define a function called `concatenate` that takes two parameters as input. (By convention, functions are also named in [snake_case](https://en.wikipedia.org/wiki/Snake_case), just like variables.)
* Inside the function, use the correct operator to concatenate those parameters
* Use a `return` statement to output the concatenated string

<!-- grade_id: define_function, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def concatenate(string_1, string_2):
    """Concatenates two strings"""
    return string_1 + string_2
### END SOLUTION
```

<!-- grade_id: define_function_test, points: 3, grade: true, solution: false -->
```python
my_concatenated_string = concatenate('Your mother was a hamster, ', 'and your father smelled of elderberries.')
print(my_concatenated_string)
### BEGIN HIDDEN TESTS
assert concatenate('Astro', 'Turf') == 'AstroTurf'
### END HIDDEN TESTS
```

Expected output:
```
Your mother was a hamster, and your father smelled of elderberries.
```

### Extension 1.3: Can you fix this error?

Run the cell below. Why doesn't it work? Can you fix your function so that is no longer vulnerable to this error? (Hint: Python has [a built-in function](https://docs.python.org/3/library/functions.html) that might help.)

```python
concatenate(7, '21')
```

## Section 2: Packages

One of the most powerful aspects of Python is the wide range of *packages* which you can import to extend the language's functionality. There are a number of useful packages included in the 'Python standard library', such as `random`, `IO` and `functools`, and there are many popular packages such as `NLTK`, `scipy` and `scikit-learn` that you can download onto your machine to add further functionality. Many of the popular additional packages are already included here in CoCalc, so for the purposes of this course you don't need to worry about downloading anything.

Once you get your skills up a bit in Python, it is surprisingly easy to develop your own packages and distribute them on the web. The ease with which packages can be created is one of the reasons Python is such a popular language today.

Importing a package is easy, using the `import` command. Execute the cell below to see how we can import the `datetime` package (part of the Python standard library), and use its `date` object to calculate today's date.

```python
import datetime

todays_date = datetime.date.today()

print(todays_date)
```

In the cell above, we imported the whole `datetime` package. But what if we just wanted to use one part of the package? It would be inconvenient to have to keep typing `datetime.date.today()` if all we wanted was `today()`. You can use the `from ... import` structure to just get the particular parts of a package you want.

```python
from nltk.tokenize import wordpunct_tokenize

my_tokens = wordpunct_tokenize('This sentence has five words.')

print(my_tokens)
```

Using `from ... import` we can avoid having to type `nltk.tokenize.wordpunct_tokenize` every time. We will be looking at 'tokenizing' text next week.

If you want even more convenience, you can use `as` to give a nickname to what you import.

```python
from matplotlib import pyplot as plt

plt.bar(['A','B','C','D','E'],[7,21,5,4,15])
plt.show()
```

We will learn about `matplotlib` and creating graphs next week. But now you have seen how to import and use functions, you can attempt...

### Assignment 1.4: Generate a random number

A simple task that is useful in lots of programming contexts is random number generation. Generating [true random numbers is complicated](https://en.wikipedia.org/wiki/Pseudorandomness). But Python allows you to easily generate 'pseudorandom' numbers, which are nearly just as good. For your first assignment of *Hacking the Humanities*, you will learn how to use the `random` package to introduce randomness into your code. The `random` package is fully documented [here](https://docs.python.org/3/library/random.html).

Your task:
* Import the `randint` function from the `random` package using `from ... import ...`. (NB: Do not use `as` for this exercise, it will confuse the grader.)
* Use the `randint` function to generate a random integer called `my_int`. (NB: the `randint()` function needs you to put a minimum and a maximum value, e.g. if you provide the values `0` and `22`, then it will generate a pseudorandom number between 0 and 22.)

<!-- grade_id: gen_rand_int, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
from random import randint
my_int = randint(0, 10)
### END SOLUTION
```

<!-- grade_id: gen_rand_int_test, points: 2, grade: true, solution: false -->
```python
# Have you created the variable correctly?
if type(my_int) == int:
    print(f'Success! You have randomly generated the integer {my_int}')
### BEGIN HIDDEN TESTS
assert randint
assert type(my_int) is int
### END HIDDEN TESTS
```

Expected output:
```
Success! You have randomly generated the integer [a random whole number]
```

## Section 3: Types and Objects

Every variable (or literal) in Python has both a **type** and a **value**. You have already encountered several types, including strings and integers. Here are some of Python's most common types:

Type | Definition | Example Value
--- | --- | ---
`str` | A string of characters | 'Ramayana' or "Ramayana"
`int` | A whole number | 891
`float` | A decimal number (a [floating point number](https://en.wikipedia.org/wiki/Floating-point_arithmetic), in computer science terms) | 75.442
`list` | An ordered collection of items | [9, 10, 10, 41]
`tuple` | An immutable\* ordered collection of items | (9, 10, 10, 41)
`dict` | A labelled collection of items | {'author':'Eliza Haywood', 'title':'Betsy Thoughtless', 'year':1751}
`set` | An unordered collection of unique items | {25,41,100028}

\* We will cover the distinction between 'mutable' and 'immutable' variables in **Week 3**.

Here is where things get a bit more complicated. Python is a so-called [object-oriented programming language](https://en.wikipedia.org/wiki/Object-oriented_programming). This means that (with a very few exceptions), nearly everything you encounter is Python is an **object** that combines *data* with *functionality*. We have already seen some free-floating functions like `print()` and `int()` and your own `concatenate()` function that you can use to do this or that using Python. But objects also have built-in functions, usually known as **methods**, which allow them to do things. When we talk about objects, we often refer to their **class** rather than their **type**, but the two words are synonymous.

One important class of object you will use throughout *Hacking the Humanities* is the `list`. It has [many useful methods](https://docs.python.org/3/tutorial/datastructures.html). Execute the cell below to see some of them in action. (**NB:** In this week's programming assignment we will use another class, `set`, which is similar to the `list` but has slightly more useful methods for our particular task.)

```python
# You define a list using square brackets
my_fruits = ['apple', 'pawpaw', 'kiwi', 'achacha', 'durian']
print(f'my_fruits when first created: {my_fruits}')

# You can add to the list using `.append()`
my_fruits.append('mango')
print(f'my_fruits after using .append(\'mango\'): {my_fruits}')

# You can retrieve the last item from the list using `.pop()`
last_fruit = my_fruits.pop()
print(f'The str \'{last_fruit}\' popped out of my_fruits')
print(f'Now my_fruits looks like this: {my_fruits}')
```

Another object that we covered in the explainer video was the **`date`** object, which can be accessed using the `datetime` package. We have already imported this package (see above), so we can just go ahead and create a new **instance** of the `date` **class** using the code below. If it doesn't work, you will need to scroll up a few cells, find the cell where we imported the `datetime` package and execute it.

```python
the_queens_birthday = datetime.date(1926, 4, 21)
```

We have created a new **instance** of the `date` **class**, which represents the Queen's birthday. Now that we have created this object, we can access many useful **attributes** and **methods**. One attribute we can access, for example, is the month the Queen was born:

```python
the_queens_birthday.month
```

A nifty *method* we can use is the `.weekday()` method. This works out what day of the week it was on the given day. What day of the week was the Queen born? (**NB:** Python counts from `0`, and considers Monday to be the first day of the week, so `0 == Monday`, `1 == Tuesday`, `2 == Wednesday` etc. etc.)

```python
the_queens_birthday.weekday()
```

If you want to have a play with the `date` class, you can find a comprehensive list of all its attributes and methods [in the official Python documentation](https://docs.python.org/3/library/datetime.html#datetime.date).

### Practice 1.5: Have a play with a `Text` object

Python packages typically contain a whole host of useful **classes** that you can use for particular applications. The `numpy` package, for instance, allows you to work with different kinds of mathematical objects, such as vectors and matrices. The `pandas` package lets you work with tables of data. The Python **Natural Language Toolkit** contains many kinds of object that are extremely useful for linguistic analysis, such as `corpora` of texts, `tokenizers` to split them up with, and different kinds of `models` for analysing them.

One of the most useful classes in the **Natural Language Toolkit** is the `Text` class. We have seen how basic types, like `str` and `int`, can be created simply by typing a string or a number. More complex classes must be created using a **constructor**. By convention, constructors are written in TitleCase. You can provide parameters to the constructor in just the same way you would to a function.

You create a new **instance** of the `Text` class by calling the constructor on a `list` of words:
```
my_text = Text(list_of_words)
```
Execute the cell below to create a `Text` out of Walt Whitman's *Leaves of Grass*.\* Then use the following cell to experiment with some of the useful methods that come with the `Text` class. Here is the [full list of the methods](https://www.nltk.org/api/nltk.html#nltk.text.Text). The `.concordance` and `.dispersion_plot` methods are particularly fun to start with.

\* We will be looking at how to import your own text data in **Week 2**.

```python
import nltk
from nltk.text import Text

leaves_of_grass_words = nltk.corpus.gutenberg.words('whitman-leaves.txt')

leaves_of_grass_text = Text(leaves_of_grass_words)

# leaves_of_grass_text is of type 'nltk.text.Text':
type(leaves_of_grass_text)
```

```python
# Have a play with some of the Text methods:
leaves_of_grass_text.concordance('myself', width=101, lines=30)
```

You can define a new class of object using a `class` statement, which works a lot like the `def` statement you use to define a function. Just like when you define a function, everything 'inside' the class has to be indented. The cell below creates a simple `class` called `KungFuFighter`. You will see that classes are a little more complex to define than functions.

The most important concept for understanding how an object works is the special `self` variable. This variable is built-in to every function and by default, every time you use a `method` of that function, `self` is provided as the first parameter. So when you type
```
leaves_of_grass_text.concordance('myself')
```
the computer interprets it as
```
leaves_of_grass_text.concordance(self, 'myself')
```
and knows that you want to apply the `.concordance()` method to *Leaves of Grass* itself.

**NB:** When you use an object method, you never need to type the word `self`——it is implied. If you look above, you will see that we typed `my_fruits.append('mango')` and NOT `my_fruits.append(self, 'mango')`.

Run the cell below to create the new class called `KungFuFighter`. The code is commented to explain to you how it works. After you have run the cell, you will be able to use the `KungFuFighter()` constructor to create new martial artists.

```python
class KungFuFighter():
    """A Kung Fu Fighter. Hee-ya!"""

    # Every class needs to have an '__init__' method. This is the method that is used to 'initialise' a new
    # instance of the class. It has double-underscores on either side of the name because it is a 'hidden'
    # method. The user will never have to type KungFuFighter.__init__(). If they want to create a new KungFuFighter,
    # they can just type `my_fighter = KungFuFighter(name=... etc.)`

    # A method is just a special kind of function, so you use a `def` statement to define it. Note how `self`
    # is the first parameter of the method, and that we indent the lines a second time to indicate we are 'inside'
    # the method:

    def __init__(self, name='Unknown fighter', hit_points=100, strength=10):

        # When we initialise a new KungFuFighter, we want to let the user choose its attributes.
        # If the user does not choose a name, hit_points or strength, then the default values will be used.
        # To store some data as an attribute, you just create a variable inside the class with a `self.` in its name.
        self.name = name
        self.hit_points = hit_points
        self.strength = strength

        # You can also add attributes that the user cannot control. By design, each KungFuFighter is alive
        # when they are created:
        self.status = 'Alive and kicking'

        # This is a general pattern: you define the attributes of a `class` in the `__init__()` method. Then you
        # define further methods below, which make use of those attributes or even change them.

        # You can access the attributes of an object in the same sort of way you access their methods. If I want to know
        # how many hit_points my KungFuFighter has, I just type `my_fighter.hit_points` or `print(my_fighter.hit_points)`.
        # Because an attribute is a variable, not a function, you do not need parentheses after `hit_points`.

    # Now we will add some more methods so our KungFuFighter can engage in combat.

    def get_kicked(self, kick_strength):
        """Get kicked by another KungFuFighter."""

        # If our KungFuFighter gets kicked, then we want their hit_points to go down according to how strong
        # the kick was.
        self.hit_points = self.hit_points - kick_strength

        # What happens when our KungFuFighter runs out of hit_points? Let's change their status from
        # 'Alive and kicking' to 'Knocked out' (this Kung Fu film is PG)
        if self.hit_points <= 0:
            self.status = 'Knocked out'
            print(f'{self.name} has been knocked out!')

        # Since we are storing the data inside the KungFuFighter, there is no need for any output, and we do not
        # need a `return` statement. Of course, sometimes you do want a method to have some sort of output. If so,
        # you can add a `return` statement as required.

    def kick(self, other_fighter):
        """Kick another KungFuFighter."""

        # If the fighter is knocked out, it can't kick!
        if self.hit_points <= 0:
            print(f'{self.name} can\'t kick, they\'ve been KO\'d!')
            # To stop the function from continuing, we can add a `return None` statement, which ends the function.
            return None

        # Print a message so that the user knows what has happened:
        print(f'{self.name} kicked {other_fighter.name}!')

        # Since other_fighter should be a KungFuFighter too, it should have the `.get_kicked()` method. We can use
        # this fact to let one KungFuFighter kick another one. We will use the strength of our own KungFuFighter
        # as the attack_strength of the kick.
        other_fighter.get_kicked(kick_strength=self.strength)

        # Again, we don't need a `return` statement, because all the data is stored inside the objects, and we do not
        # need any output.
```

Once a class has been defined, you can use the `help()` function to find out more about it. This works for all classes, not just the ones you have defined yourself:

```python
help(KungFuFighter)
```

### Practice 1.6: Create and use a KungFuFighter instance

Execute the below code cell to create two KungFuFighters.

```python
fighter_one = KungFuFighter(name="Fighter One", hit_points=50, strength=30)
fighter_two = KungFuFighter(name="Fighter Two", hit_points=180, strength=15)
```

Now you can use the below cell to make `fighter_two` attack `fighter_one`. How many times do you need to execute the cell before `fighter_one` is KO'd?

```python
fighter_two.kick(fighter_one)
```

Now have a play around with the two code cells. If you re-execute the first cell, the fighters will be re-created with their original attibutes. You can also alter their `hit_points` and `strength` as you please. You can then change the code in the second cell to change the course of the battle. If you add additional lines of code you can change who attacks whom, and how many times.

### Assignment 1.7 Everybody was Kung Fu Fight-ing!

For this assignment, you are going to create *two* Kung Fu fighters and battle them against one another.

Your task:
* Create two fighters, `lady_sun` and `cao_cao`. `lady_sun` should have the name "[Sun Shangxiang](https://en.wikipedia.org/wiki/Lady_Sun)", have `150` hit points, and have a strength of `30`. `cao_cao` should be called "[Cao Cao]("https://en.wikipedia.org/wiki/Cao_Cao")", have `90` hit points and have a strength of `50`.
* It's the battle of [Red Cliffs](https://en.wikipedia.org/wiki/Battle_of_Red_Cliffs). Lady Sun attacks Cao Cao until he is knocked out. (**NB:** You need to use `lady_sun`'s' the `.kick()` method to make this work. To complete the assignment correctly, you need to make the correct number of kicks occur from just *one* execution of the cell.)

<!-- grade_id: kung_fu_fighting, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
lady_sun = KungFuFighter(name="Sun Shangxiang", hit_points=150, strength=30)
cao_cao = KungFuFighter(name="Cao Cao", hit_points=90, strength=50)
lady_sun.kick(cao_cao)
lady_sun.kick(cao_cao)
lady_sun.kick(cao_cao)
### END SOLUTION
```

<!-- grade_id: kung_fu_fighting_tests, points: 5, grade: true, solution: false -->
```python
# Is lady_sun of the right type?
print("Lady Sun's type: " + str(type(lady_sun)))

# Did she KO Cao Cao?
print("Cao Cao's status: " + cao_cao.status)
### BEGIN HIDDEN TESTS
assert type(lady_sun) is KungFuFighter
assert type(cao_cao) is KungFuFighter
assert lady_sun.name == "Sun Shangxiang"
assert lady_sun.strength == 30
assert cao_cao.name == "Cao Cao"
assert cao_cao.strength == 50
assert cao_cao.hit_points <= 0 and cao_cao.status == 'Knocked out'
### END HIDDEN TESTS
```

Expected output:
```
Lady Sun's type: <class '__main__.KungFuFighter'>
Cao Cao's status: Knocked out
```

## Section 4: Creating a `ShoppingList` Object

With the knowledge you have of Python's basic structures, you are now ready for your first big assignment: to define your own **object class** and use it to create your own **instances**.

Based on what we covered in the lecture, you are going to build a class for making shopping lists in Python. You will first define a ne **class** called `ShoppingList`. You will then use this **class** to create several different shopping lists for different shops you need to visit. Each of these particular shopping lists will be an **instance** of the `ShoppingList` **class**.

To be useful, the `ShoppingList` class needs to be able to do four things:
* Store information about which shop you are going to (we will use an **attribute** to do this)
* Store a list of items you need to buy inside it (we will use an **attribute** to do this)
* Allow you to add new items to the list (we will use a **method** to do this)
* Allow you to remove items from the list once you have bought them (we will use another **method** to do this)

You should be able to see how the `ShoppingList` will be similar to `KungFuFighter`. `KungFuFighter` also stored some data (the `name`, `hit_points` and `strength` of the fighter) and had some useful methods (`.kick()` and `.get_kicked()`).

This Section is in two parts. First you need to define `ShoppingList` using a `class` statement. Then in the following section you need to create two **instances** of this class for two shops that you are going to visit.

### Assignment 1.8: Define the `ShoppingList` Class

For this assignment, you need to define the class. I have already provided the `class` statement and docstring to get you going. To complete the assignment, you need to do three things:
1. Write an `__init__` method that defines the class's attributes. The `ShoppingList` class needs two attributes: `self.shop`, the name of the shop that the list is for, and `self.items`, the items that the user wants to buy at the shop. `self.items` should be of the `set` type. **Remember:** Here you are defining the **class**, not a particular **instance** of a shopping list. Your `__init__` function will take information from the user and save it inside the object. I have provided some guidance in the code cell below.
2. Write an `add_item` method that allows the user to add an item to their shopping list. If you correctly defined `self.items` as a `set` in step 1, then you will be able to make use of `self.items`'s built-in `.add()` method.
3. Write a `buy_item` method that removes an item from the shopping list once the user has bought it. If you have correctly defined `self.items` as a `set` in step 1, then you will be able to make use of `self.items`'s `.remove()` method.

**Hints:**
* I show how to complete this assignment in detail in Explainer Video 1 - 3, starting at about 10 minutes.
* Make sure you keep track of your indentation. Each `def` statement (for `__init__`, `add_item` and `buy_item`) should be indented 4 spaces. Everything *inside* any of those `def` statements should be indented 4 spaces more
* None of the functions you define inside `ShoppingList` need a `return` statement, as they are all object **methods** that manipulate the object itself rather than returning an output to the user

<!-- grade_id: define_shopping_list_class, points: 0, grade: false, solution: true -->
```python
class ShoppingList(object):
    """A class for creating and using shopping lists."""

    # First you need to define the `__init__` method. This is the hidden method that Python
    # will use to initialise a new shopping list each time the user creates a new instance of
    # the ShoppingList class.
    # You can use the `__init__` method to define some attributes for the object.
    # All attributes should have `self.` put in front of them. The `self.` tells Python that you
    # want the relevant data stored inside *this* object.
    # We need two attributes for this class: `shop` and `items`.

    # Step 1: Use a `def` statement to define a function calld `__init__`. This function should
    # have two parameters: `self` and `shop`. `self` is a hidden parameter that helps the class
    # work correctly. `shop` will be the name of a shop, inputted by the user.
    def __init__(self, shop):
        # Step 2: Inside the `__init__` function, define the first attribute, `shop`. This attribute
        # should store the shop name inputted by the user. Remember, to define an attribute you can use
        # the template self.attribute = parameter. In this case, the attribute should be called `shop`
        # and the paramter happens also to be called `shop`.
        self.shop = shop
        # Step 3: Still inside the `__init__` function, define the second attribute, `items`. This
        # attribute will be empty of data, because we have not asked the user to input any items
        # when they first create the ShoppingList object. (The only thing we ask is for them to 
        # enter a `shop` paraemter.) BUT - we still need to define the `items` attribute so it is
        # ready to store the user's data when they input it. AND we want the `items` attribute to be
        # of the type `set`. If you want to create an empty attribute of a particular type, you can use
        # the template `self.attribute = type_name()` (e.g. self.something = list() would create an
        # attribute called `something` of type `list`). In this case,the attribute should be `items` and
        # the type_name is `set`.
        self.items = set()

    # Step 4: Use a new `def` statement to define the `add_item` method. This method should have two
    # parameters: `self` and then `item`. Add the following docstring to the method: """Add item to my shopping list"""
    def add_item(self, item):
        """Add item to my shopping list"""
        # Step 5: Write the code that will add the item inputted by the user to the items attribute.
        # If you have correctly made self.items a `set` object, then you will be able to use its built-in
        # `.add()` method. Use the following template for your code: self.attribute.method(parameter).
        # The attribute is `items`, the method is `add`, and the parameter is `item`.
        self.items.add(item)

    # Step 6: Use a new `def` statement to define the `buy_item` method. This method should look very
    # similar to the `add_item` method. It should also have the parameters `self` and `item`. But the
    # docstring should be different: """Tick off item once I've bought it"""
    def buy_item(self, item):
        """Tick off item once I've bought it"""
        # Step 6: Write the code that will remove the item from the list once you've bought it.
        # This code will look very similar to the code for `add_item`, but with one crucial
        # difference. This time you shoudl use the `.remove()` method of `self.items`, rather than
        # the `.add()` method.
        self.items.remove(item)
```

<!-- grade_id: 361e97, points: 10, grade: true, solution: false -->
```python
help(ShoppingList)
### BEGIN HIDDEN TESTS
assert hasattr(ShoppingList, "add_item")
assert ShoppingList.add_item.__doc__ == "Add item to my shopping list"
assert hasattr(ShoppingList, "buy_item")
assert ShoppingList.buy_item.__doc__ == "Tick off item once I've bought it"
test_list = ShoppingList('just a random string')
assert hasattr(test_list, "shop")
assert test_list.shop == 'just a random string'
assert hasattr(test_list, "items")
assert isinstance(test_list.items, set)
test_list.add_item("klingons")
assert "klingons" in test_list.items
assert len(test_list.items) == 1
test_list.buy_item("klingons")
assert len(test_list.items) == 0
### END HIDDEN TESTS
```

If you have defined the `ShoppingList` class correctly, then the output of the above cell should look like so:
```Help on class ShoppingList in module __main__:

class ShoppingList(builtins.object)
 |  ShoppingList(shop)
 |  
 |  A class for creating and using shopping lists.
 |  
 |  Methods defined here:
 |  
 |  __init__(self, shop)
 |      Initialize self.  See help(type(self)) for accurate signature.
 |  
 |  add_item(self, item)
 |      Add item to my shopping list
 |  
 |  buy_item(self, item)
 |      Tick off item once I've bought it
 |  
 |  ----------------------------------------------------------------------
 |  Data descriptors defined here:
 |  
 |  __dict__
 |      dictionary for instance variables (if defined)
 |  
 |  __weakref__
 |      list of weak references to the object (if defined)
```
Some things to check:
* Are all the required methods there? (`__init__`, `add_item` etc.)
* Are all the docstrings correct?
* Does each method have the correct parameters? (`self`, `shop` etc.)

### Assignment 1.9: Create some `ShoppingList`s!

Now you have defined the `ShoppingList` **class**, you can create some **instances** of the class to do some shopping.

You can create a `ShoppingList` using the below template. This template would create a new variable `grocery_list`, which would be for the `'Groceries'`. It would have two items on the list, `potatoes` and `cheese`. To see what items are on the list, we can use `print()`.
```
grocery_list = ShoppingList(shop='Groceries')
grocery_list.add_item('potatoes')
grocery_list.add_item('cheese')
print(grocery_list.items)

Expected output:
potatoes
cheese
```

To complete this assignment, you are going to create two `ShoppingList`s, and make use of their `add_item` and `buy_item` methods.

This is what you need to do:
1. Create a shopping list for the bakery. Use `bakery_list` for the variable name, and make `'Bakery'` the `shop` attribute. (1 line of code)
2. Add two items to `bakery_list`: `'rolls'` and `'sponge cake'` (2 lines)
3. Create another shopping list, this time for the homewares store. Use `homewares_list` for the variable name, and `'Lakeland'` for the `shop` attribute. (1 line)
4. Add one item to `homewares_list`: `'pressure cooker'` (1 line)
5. You are at the bakery. They have run out of rolls, but they still have sponge cake. Use the `.buy_item()` method to buy the `'sponge cake'`. (1 line)

<!-- grade_id: create_shopping_lists, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
bakery_list = ShoppingList(shop='Bakery')
bakery_list.add_item('rolls')
bakery_list.add_item('sponge cake')
homewares_list = ShoppingList(shop='Lakeland')
homewares_list.add_item('pressure cooker')
bakery_list.buy_item('sponge cake')
### END SOLUTION
```

<!-- grade_id: 9ee4d9, points: 5, grade: true, solution: false -->
```python
print("What's left to buy?")
print(f"At the {bakery_list.shop} I need to buy: {bakery_list.items}")
print(f"At {homewares_list.shop} I need to buy: {homewares_list.items}")
### BEGIN HIDDEN TESTS
assert isinstance(bakery_list, ShoppingList)
assert isinstance(homewares_list, ShoppingList)
assert 'rolls' in bakery_list.items
assert 'sponge cake' not in bakery_list.items
assert 'pressure cooker' in homewares_list.items
assert 'Bakery' == bakery_list.shop
assert 'Lakeland' == homewares_list.shop
### END HIDDEN TESTS
```

**Expected output:**
```
What's left to buy?
At the Bakery I need to buy: {'rolls'} <-- 'sponge cake' should have disappeared
At Lakeland I need to buy: {'pressure cooker'}
```

## Conclusion

Congratulations! You've made it to the end of Week 1, and learned some serious programming skills along the way. You now know all the fundamental concepts of object-orientd programming in Python: variables, operators, functions, objects, classes, attributes and methods. You can create your own unique software by defining your own custom functions and--even cooler--by defining your own unique object classes that let you package together data and functionality.

In Week 2 we will start to see how you can apply your new skills to literary analysis.

### Extension 1.20

If you want to challenge yourself, you could try creating a new version of your `ShoppingList` object with improved functionality. You can do this easily using subclassing. I have provided the necessary code below. When you create a subclass, you create a new object which has all the same attribute and methods as its parent object by default. Then you can add a new method or two, or overwrite an existing method as you wish.

Some things you could try:
* Redefine the `__init__` method so that the user can add items to the shopping list when it is first initialised
* Add a new method that displays which items are in the list, so that the user doesn't need to use the `print()` function
* **Advanced:** Change the `.add_item` and `.buy_item` methods so that users can add or buy multiple items at once. There is a clever trick that will let you do this by simply adding a single character to the `items` parameter in the function definition, and then below when you use `.add()` or `.remove()`. The trick is called unpacking. You can [read about it here](https://docs.python.org/3/tutorial/controlflow.html?highlight=unpacking#unpacking-argument-lists).

```python
class ImprovedShoppingList(ShoppingList):
    """New and improved ShoppingList!"""

    # Make any changes you like below...
```